Fitness drivers of division of labor in vertebrates

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #2 (Public review):

Summary:

This paper formulates an individual-based model to understand the evolution of division of labor in vertebrates. The model considers a population subdivided in groups, each group has a single asexually-reproducing breeder, other group members (subordinates) can perform two types of tasks called "work" or "defense", individuals have different ages, individuals can disperse between groups, each individual has a dominance rank that increases with age, and upon death of the breeder a new breeder is chosen among group members depending on their dominance. "Workers" pay a reproduction cost by having their dominance decreased, and "defenders" pay a survival cost. Every group member receives a survival benefit with increasing group size. There are 6 genetic traits, each controlled by a single locus, that control propensities to help and disperse, and how task choice and dispersal relate to dominance. To study the effect of group augmentation without kin selection, the authors cross-foster individuals to eliminate relatedness. The paper allows for the evolution of the 6 genetic traits under some different parameter values to study the conditions under which division of labor evolves, defined as the occurrence of different subordinates performing "work" and "defense" tasks. The authors envision the model as one of vertebrate division of labor.

The main conclusion of the paper is that group augmentation is the primary factor causing the evolution of vertebrate division of labor, rather than kin selection. This conclusion is drawn because, for the parameter values considered, when the benefit of group augmentation is set to zero, no division of labor evolves and all subordinates perform "work" tasks but no "defense" tasks.

Strengths:

The model incorporates various biologically realistic details, including the possibility to evolve age polytheism where individuals switch from "work" to "defense" tasks as they age or vice versa, as well as the possibility of comparing the action of group augmentation alone with that of kin selection alone.

Weaknesses:

The model and its analysis are limited, which in my view makes the results insufficient to reach the main conclusion that group augmentation and not kin selection is the primary cause of the evolution of vertebrate division of labor. There are several reasons.

(1) First, although the main claim that group augmentation drives the evolution of division of labor in vertebrates, the model is rather conceptual in that it doesn't use quantitative empirical data that applies to all/most vertebrates and vertebrates only. So, I think the approach has a conceptual reach rather than being able to achieve such a conclusion about a real taxon.

We appreciate the reviewer’s point that our model does not incorporate quantitative empirical data across vertebrate taxa. This is indeed a limitation and reflects the current lack of fine-scale datasets on task division, the influence of life-history traits, and the fitness consequences of different cooperative activities in vertebrates. One of our aims, however, is precisely to stimulate such empirical work by highlighting the value of examining division of labor in species inhabiting harsh environments, considering age/size/dominance structure when evaluating variation in cooperative activities, and incorporating defense behaviors more consistently into analyses of helping, especially since defenders are often overlooked relative to the classic helpers-at-the-nest that provision offspring. The model therefore remains directly relevant to vertebrate systems because it departs from insect-inspired approaches that focus on fitness outcomes based solely in maximizing colony productivity. Instead, it incorporates direct fitness benefits to group members, an essential feature of vertebrate cooperative breeding and of other systems with fertile “workers,” as we clarified in the discussion.

(2) Second, I think that the model strongly restricts the possibility that kin selection is relevant. The two tasks considered essentially differ only by whether they are costly for reproduction or survival. "Work" tasks are those costly for reproduction and "defense" tasks are those costly for survival. The two tasks provide the same benefits for reproduction (eqs. 4, 5) and survival (through group augmentation, eq. 3.1). So, whether one, the other, or both helper types evolve presumably only depends on which task is less costly, not really on which benefits it provides. As the two tasks give the same benefits, there is no possibility that the two tasks act synergistically, where performing one task increases a benefit (e.g., increasing someone's survival) that is going to be compounded by someone else performing the other task (e.g., increasing that someone's reproduction). So, there is very little scope for kin selection to cause the evolution of labor in this model. Note synergy between tasks is not something unusual in division of labor models, but is in fact a basic element in them, so excluding it from the start in the model and then making general claims about division of labor is unwarranted. In their reply, the authors point out that they only consider fertility benefits as this, according to them, is what happens in cooperative breeders with alloparental care; however, alloparental care entails that workers can increase other's survival *without group augmentation*, such as via workers feeding young or defenders reducing predator-caused mortality, as a mentioned in my previous review but these potentially kin-selected benefits are not allowed here.

We understand the reviewer’s concern that our model restricts the scope for kin-selected benefits by not including task-specific synergy effects—specifically, help that directly increases the survival of group members (e.g., load-lightening via feeding young, or predator defense that reduces mortality of breeders or offspring independently of group augmentation). We agree that such effects can occur in some cooperative breeders, and that they can, in principle, generate indirect fitness benefits. However, even when helpers increase the survival of breeders or reduce parental investment per offspring, these effects generally translate into higher breeder productivity—either via increased fecundity, increased survival to the next breeding attempt, or increased investment in subsequent broods. Thus, although we treat benefits in terms of enhanced breeder productivity, this formulation implicitly captures a range of help-related effects that ultimately improve the reproductive output of the breeders, including those mediated through increased survival. For this reason, we believe that the model remains relevant for vertebrate systems despite not representing each pathway separately.

(3) Third, the parameter space is understandably little explored. This is necessarily an issue when trying to make general claims from an individual-based model where only a very narrow parameter region of a necessarily particular model can be feasibly explored. As in this model the two tasks ultimately only differ by their costs, the parameter values specifying their costs should be varied to determine their effects. In the main results, the model sets a very low survival cost for work (yh=0.1) and a very high survival cost for defense (xh=3), the latter of which can be compensated by the benefit of group augmentation (xn=3). Some limited variation of xh and xn is explored, always for very high values, effectively making defense unevolvable except if there is group augmentation. In this revision, additional runs have been included varying yh and keeping xh and xn constant (Fig. S6), so without addressing my comment as xn remains very high. Consequently, the main conclusion that "division of labor" needs group augmentation seems essentially enforced by the limited parameter exploration, in addition to the second reason above.

As we have explained in previous revisions, the costs associated with work and defense are not directly comparable because they affect different fitness components: work costs reduce dominance, whereas defense costs reduce survival. Whether a particular cost is “high” or “low” can only be evaluated by examining the evolved reaction norms and identifying the ranges over which these norms change. For this reason, we focused on parameter ranges that actually generate shifts in reaction norms rather than presenting large regions of parameter space where nothing changes.

We also reiterate that we did in fact explore broader parameter ranges than those shown in the main text. Additional analyses, including those specifically designed to identify conditions under which division of labor evolves under kin selection alone, are provided in the Supplementary Material. Specifically, Figure S1 addresses the point raised by the “need” of group augmentation benefits for defense to evolve, by increasing the baseline survival x0.

We now include one additional figure in the Supplementary Material with a lower value for the benefit of group size (xn = 1 instead of xn = 3), and we extended the range of xh to include lower values (xh = 1). As we can see in Figure S7 and Table S8, group augmentation benefits are still the primary reason for individuals to group (see dispersal values). For low benefits of group augmentation, defense evolves in harsh environments in the absence of kin selection, and in benign environments when both direct and indirect fitness benefits take place. We have also now expanded the results section to include these last results. Note that we also checked even lower values for xh under the only kin selection implementation, with results being qualitatively similar, but chose not to include them in the manuscript since it is already a very long Supplementary Material. Here are the averages for two examples with xh = 0.1 and when we promote division of labor:

In short, the conclusion that division of labor requires group augmentation is not an artifact of limited parameter exploration. It arises because kin selection alone favors division of labor only under highly restrictive parameter combinations, whereas including direct fitness benefits substantially expands the conditions under which division of labor evolves. This pattern is consistent across the full set of parameter combinations we examined.

(4) Fourth, my view is that what is called "division of labor" here is an overinterpretation. When the two helper types evolve, what exists in the model is some individuals that do reproduction-costly tasks (so-called "work") and survival-costly tasks (so-called "defense"). However, there are really no two tasks that are being completed, in the sense that completing both tasks (e.g., work and defense) is not necessary to achieve a goal (e.g., reproduction). In this model there is only one task (reproduction, equation 4,5) to which both helper types contribute equally and so one task doesn't need to be completed if completing the other task compensates for it; instead, it seems more fitting to say that there are two types of helpers, one that pays a fertility cost and another one a survival cost, for doing the same task. So, this model does not actually consider division of labor but the evolution of different helper types where both helper types are just as good at doing the single task but perhaps do it differently and so pay different types of costs. In this revision, the authors introduced a modified model where "work" and "defense" must be performed to a similar extent. Although I appreciate their effort, this model modification is rather unnatural and forces the evolution of different helper types if any help is to evolve.

In previous models of division of labor in eusocial insects, the implicit benefit is also colony-level productivity (see Beshers & Fewell, 2001, for a review of division of labor in insects). Even in humans, division of labor functions as a means to increase efficiency toward achieving a shared goal. Our model adopts this same interpretation, as outlined in the Introduction, but extends it by considering that different tasks may impose different fitness costs, an aspect that has been largely overlooked in the existing literature. It is precisely because fitness outcomes are not fully shared among group members in vertebrates that distinguishing these cost structures matters. Unlike eusocial insects with sterile workers, vertebrate helpers can obtain direct fitness benefits, and the model explicitly accounts for these direct benefits—something absent from most insect-inspired approaches even when direct fitness benefits can also arise in some of those systems. Thus, our framework is not simply evolving “two types of helpers doing the same task,” but instead evolving specialization in different cooperative roles that carry different fitness consequences. It is therefore suitable for our model to treat contributions to breeder productivity as a common currency, while allowing individuals to specialize in different cost-distinct forms of help.

Finally, regarding synergy: with the extension introduced in the previous revision, we now incorporate the requirement that multiple forms of help must be performed for the group to achieve maximal reproductive output. This directly addressed the reviewer’s concern about synergistic dependencies between tasks and aligns our framework with the kinds of complementarity highlighted in other models of division of labor.

In summary, the structure of the model is consistent with both the theoretical literature on division of labor and the biological realities of vertebrate cooperative systems. We believe it is important for future models to explicitly consider the different fitness benefits and costs associated with distinct cooperative behaviors, and hope that our framework encourages more targeted empirical research on division of labor in vertebrates (e.g. inclusion of data on defense, life-history traits and environmental challenges) to better inform future modelling efforts.

I should end by saying that these comments don't aim to discourage the authors, who have worked hard to put together a worthwhile model and have patiently attended to my reviews. My hope is that these comments can be helpful to build upon what has been done to address the question posed.

We appreciate the reviewer’s thoughtful and constructive comments, as well as the time invested in evaluating our work. These insights have greatly helped us improve the clarity and overall quality of the manuscript. We hope that the revisions and additional clarifications we have provided adequately address all remaining concerns.

Changes from the earlier version to the revised version

Abstract

Although division of labor as a means to increase productivity is a common feature in animal social groups, most previous studies have focused almost exclusively on eusocial insects with extreme task partitioning. Empirical evidence of division of labor in vertebrates is scarce, largely because we lack a theoretical framework to explore the conditions under which division of labor is likely to evolve in cooperatively breeding systems where helpers remain capable of breeding throughout their lifetime. By explicitly considering alternative helping tasks with varying fitness costs, we model how individual decisions on task specialization may influence the emergence of division of labor under both direct and indirect fitness benefits. Surprisingly, we find that direct survival benefits of living in larger groups are the primary force driving the evolution of cooperation to enhance group productivity, and that indirect fitness benefits derived from related group members are only a non-essential facilitator of more stable forms of division of labor in cooperative breeders. In addition, we find that division of labor in vertebrates is favored by harsh environments. Ultimately, our model not only makes key predictions that are consistent with existing empirical data, but also proposes novel avenues for new empirical work in vertebrate and invertebrate systems alike.

Introduction

Division of labor as a means to boost productivity holds significance not only in human societies, but also those of other animals [1–5]. Nearly 250 years ago, the economist Adam Smith [2] proposed that division of labor via task specialization within human societies not only improves individual efficiency and group productivity, but also reduces the costs associated with switching tasks. Similarly, in nonhuman animal societies, division of labor to increase group reproductive output (analogous to economic gain) is characterized by group members specializing in particular tasks, either temporarily or permanently [6–8]. To date, most theoretical and empirical work on the selective pressures favoring division of labor via task specialization in animal societies has focused on “classically” eusocial insect societies characterized by permanently sterile workers that are never able reproduce [8–13]. Since workers in eusocial societies can, in most cases, only obtain fitness benefits indirectly through kin selection by maximizing their colony reproductive output, selection can occur at the colony level, typically resulting in very low conflict of interest between group members. However, complex societies occur outside of eusocial insects, such as in many birds and mammals, including humans. Because sterile workers are absent in vertebrates and other cooperative breeding societies, conflict can be much higher than in most eusocial insect societies. The few studies of division of labor outside of eusocial insects have primarily focused on reproductive division of labor (i.e., reproductive skew), despite that fact that temporal or permanent non-reproductive organisms can engage in an array of different helping tasks [3, 14, 15]. Given the scarcity of studies examining the selective advantages of division of labor beyond eusocial insects (but see [16]), our understanding of the evolutionary pressures of division of labor is largely confined to instances where task specialization is linked to sterility and where conflict of interest is generally low [17].

In cooperatively breeding vertebrates, which typically live in societies with more equitable sharing of reproduction than eusocial insects (i.e., low reproductive skew), low or mixed kinship, and often have unrelated helpers, conflict of interest can be high and direct fitness benefits may be more relevant in the evolution of division of labor than indirect benefits [18–20]. In fact, shared fitness incentives to increase group productivity are not limited to eusocial insects with high within-group relatedness and sterile workers. In cases where group size positively correlates with an increase in the survival or reproduction of group members, division of labor via task specialization may also evolve under direct fitness benefits to subordinate members of a group (i.e., group augmentation [21, 22]). Although a great deal of recent empirical work highlights the importance of direct benefits in the evolution of cooperative breeding behavior in vertebrates [23–26], we still lack an understanding of the joint influence of direct and indirect fitness benefits in the evolution of division of labor.

Understanding how direct and indirect benefits interact is particularly important in systems where individuals may differentially bear the fitness costs of cooperation. Most previous models of division of labor via task specialization considered all helping tasks under the same fitness umbrella, which is a good approximation in eusocial insect societies because colony members share fitness costs and benefits. In vertebrates and social insects with fertile workers, however, helpers may obtain fitness benefits directly via reproduction, and thus, selection occurs more strongly at the individual level than it does in most eusocial insects. Group members in vertebrate societies may then contribute to distinct helping tasks differently according to individual life history strategies [27–29] or differences in the magnitude or type of fitness costs and benefits that they receive from cooperative tasks [30]. One of the key distinctions between fitness costs is whether a task impacts more immediate survival or direct reproduction. For instance, defensive behaviors such as predator mobbing may have a high risk of injury and death [31–33], whereas other activities such as allofeeding or nest/territory maintenance may affect body condition owing to lower feeding rates and time and energy investment, consequently affecting individual reproductive output [34–36]. Therefore, changes in life histories, such as those linked to the likelihood of becoming a breeder, may select for task specialization according to the relative direct fitness cost associated with different helping tasks.

In addition to life history differences, demographic, social, and ecological factors that affect help workload and the prospect of helpers becoming breeders are likely to have a strong impact on task allocation because they alter the probability of obtaining direct fitness benefits [37, 38]. For example, age is known to be an important factor influencing helping workload in vertebrates [39, 40], is often correlated with differences in other life-history traits [36, 41–43], and may be a proxy for other processes involved in the evolution of help [44, 45]. Moreover, age-dependent task allocation (i.e., temporal polyethism) is common in eusocial insects that do not show morphological differentiation [46]. Social and ecological factors are also thought to influence the emergence of division of labor in eusocial insects [47–49]. Indeed, high relatedness may reduce conflict of interest between group members and promote the evolution of division of labor [50–54]. Similarly, in harsh environments where group size strongly influences survival [55–58], division of labor may be favored to more efficiently increase group productivity.

To disentangle the role of direct versus indirect fitness benefits in the evolution of division of labor in cooperative breeding societies with totipotent helpers, we created an individual-based model in which we varied the influence of different evolutionary forces to understand under which social and ecological circumstances division of labor may arise to maximize the inclusive fitness of individuals in a group. We used distinct tasks with different fitness costs to maximize the group reproductive output when they were performed to a similar extent. Individuals in our model were allowed to evolve different behavioral responses or reaction norms in task specialization, depending on their probability of becoming breeders. We compared the results of this model to those of a benchmark model in which the reproductive outcome was independent of the type of task performed by group members. We present the results for different environmental conditions in the presence or absence of relatedness (indirect) or group size (direct) benefits. Ultimately, this work demonstrates that the direct fitness benefits of group living attained in increasingly harsh environments are key to the evolution of division of labor in vertebrates.

Results

Rationale

Our model consists of a population characterized by overlapping generations residing in a habitat featuring a fixed finite number of breeding territories. These territories are monopolized by groups composed of a single breeder and a variable number of subordinate helpers, the number of which is shaped by population dynamics, with all group members capable of reproducing during their lifetime. Individuals in the population fall into one of three categories: (1) breeders that monopolize reproduction within a group; (2) subordinates that may perform different helping tasks, either in their natal group or in an outside group to which they dispersed; and (3) floaters that disperse from their natal group but remain not part of a group.

We assume complete reproductive skew such that subordinate group members are prevented from reproducing at a given timestep and may only help raise the offspring of the dominant breeders. Dispersers cannot reproduce without acquiring a territory (denoted here as floaters). Helpers are not sterile and may inherit the breeding position when the breeder dies. When this occurs, helpers and dispersers from other groups compete for the breeding position and win with a probability proportional to their dominance value (Figure 1; Table 1). We use the term dominance value to designate the competitiveness of an individual compared to other candidates in becoming a breeder, regardless of group membership. Dominance value increases as a function of age, serving as a proxy for resource holding potential (RHP), and decreases as a function of help provided, reflecting costs to body condition from performing working tasks (Eq. 2).

Diagram of the scheduling executed per breeding cycle.

(1) A breeder reproduces. Its productivity depends on the cumulative level of brood care provided by the group during the previous breeding cycle. Maximum productivity is achieved when different helping tasks are performed to a similar extent. (2) Subordinates may disperse to become floaters, or they may stay in the group and help. Dispersers/floaters may join a random group to become subordinates. (3) Subordinates in the group (both natal and immigrant individuals) either work to provision to the breeder’s offspring or display defensive forms of help. (4) Individuals survive contingent on group-living benefits and dispersal costs, as well as the cost of defensive activities. (5) If a breeder dies, helpers in the group and a sample of floaters compete for the breeding position. Individuals still alive ascend one age class, and the cycle starts all over (i.e., next breeding cycle).

Subordinate helpers in our model may engage in one of two distinct tasks with differing fitness costs: (1) defensive tasks (i.e., defense of the dominant breeders’ offspring at cost to their own survival); or (2) work tasks (i.e., tasks with an associated cost to their body condition and consequently to their competitiveness to become breeders in the future). Subordinate helpers can then evolve different workloads and exhibit task specialization that is either fixed throughout their lifetime or varies with their dominance value. The parameters xh and yh (Table 1) reflect how much the subordinates’ survival and dominance values decrease with increasing investment in defense or work tasks, respectively. Although we do not assume an inherent difference in body condition between individuals when born, variance in age and helping propensities can create differences in the fitness costs linked to task performance.

We assume net survival benefits from living in large groups (i.e., group augmentation), which occurs in many social species [56, 57, 59, 60]. Subordinate members of a group may therefore gain both direct and indirect fitness benefits from helping to increase the productivity of breeders, if doing so results in an increase in group size. In the model, the benefit of alloparental care to the productivity of the breeder is greatest when group members provide overall similar workloads for both defensive and work tasks, since both are often needed for the survival of the offspring [61, 62]. For example, in many cooperatively breeding birds, the primary reasons that breeders fail to produce offspring are (1) starvation, which is mitigated by the feeding of offspring, here considered as a work task, and (2) nest depredation, which is countered by defensive behavior [61, 63]. Consequently, both types of tasks are often necessary for successful reproduction, and focusing solely on one while neglecting the other is likely to result in lower reproductive success than if both tasks are performed by helpers within the group. To compare between the helper’s preferred task and the task needed to maximize group productivity, we created a benchmark model in which the breeder’s productivity is not constrained by the type of help provided. Higher productivity translates to larger group sizes, which enhances the survival prospects of all group members. Hence, subordinates have an incentive to develop division of labor. Subordinates that remain in their natal group may also gain improved fitness by indirectly passing on their genes through kin. Both group size and relatedness in our model are emergent properties of a series of demographic processes, including mortality, dispersal, and help-dependent breeder fertility.

To distinguish between the effects of group augmentation and kin selection on inclusive fitness, we created two additional models. In the first model, we removed the benefits of group size (group augmentation) by setting the parameter xn = 0 (Table 1). In the second model, we removed the effect of kin selection by simulating a cross-foster manipulation in which offspring just born are interchanged between different groups (see Methods for additional details). To assess how habitat quality affects task specialization and the emergence of division of labor, we systematically varied habitat quality across simulations. Harsh environments are those that have high mortality (i.e., m = 0.3; Table 1), which also raises the turnover of the breeding positions. Lower values of m reflect an increased probability of individuals surviving another breeding cycle and, therefore, the chances of habitat saturation. The steps in the breeding cycle of the model are illustrated in Figure 1.

Direct versus indirect fitness benefits

We first examined how both direct fitness benefits from group living and indirect fitness benefits from aiding kin influence the evolution of division of labor through task specialization across environments of varying quality. Our results suggest that voluntary division of labor involving tasks with different fitness costs is more likely to emerge initially because of direct fitness benefits (Figure 2; GA), but a combination of direct and indirect fitness benefits leads to higher rates and more stable forms of division of labor (Figure 2; KS + GA). It is important to note that, as depicted in Figure 2, intermediate values of task specialization indicate in all cases age/dominance-mediated task specialization (γR ≠ 0; Table 1) and never a lack of specialization (γR = 0; Table 1); further details are shown in the next section: “The role of dominance in task specialization”.

Effect of environmental quality on alloparental care and division of labor.

Earlier version: The evolutionary equilibria for phenotypic levels of helping and task specialization are shown at three different levels of environmental quality, ranging from benign (m=0.1) to harsh (m=0.3), and for three different levels of cost of help on survival (light blue, xh=3; blue, xh=5; and dark blue, xh=7), across 20 replicas. The vertical axis represents the expressed task allocation between a defensive task with a cost to survival versus a work task with a cost to their dominance rank. The optimum breeder productivity per unit of help provided was either when both tasks were performed to a similar extent, potentially selecting for division of labor (▴) or when no restrictions were introduced to the task performed by the group members (〇). In each environment, additional details are given on the selective forces that play a role in the evolution of help and task specialization: help can only evolve by kin selection (KS), group augmentation (GA), or both (KS + GA). Additional details are provided in Table S1. For variation in yhvalues, see Figure S6. All input parameter values are described in Table 1.

Revised version: The evolutionary equilibria for phenotypic levels of helping and task specialization are shown at three different levels of environmental quality, ranging from benign (m=0.1) to harsh (m=0.3), and for three different levels of cost of help on survival (light blue, xh=3; blue, xh=5; and dark blue, xh=7), across 20 replicas. The vertical axis represents the expressed task allocation between a defensive task with a cost to survival versus a work task with a cost to their dominance rank. The optimum breeder productivity per unit of help provided was either when both tasks were performed to a similar extent, potentially selecting for division of labor (▴) or when no restrictions were introduced to the task performed by the group members (〇). In each environment, additional details are given on the selective forces that play a role in the evolution of help and task specialization: help can only evolve by kin selection (KS), group augmentation (GA), or both (KS + GA). Additional details are provided in Table S1. For variation in yh values, see Figure S6. All input parameter values are described in Table 1.

Earlier version: In contrast to expectations from models of eusocial insects, our vertebrate model does not readily evolve division of labor when only kin selection is considered (see the section “Kin selection and the evolution of division of labor” and Figures S1-S4 in the Supplemental Material for more details on further exploration on conditions that may favor the evolution of division of labor under only kin selection). Specifically, forms of help that impact survival rarely evolve under any environmental condition when only kin selection occurs (Figure 2; KS). This occurs because with kin competition for the breeding position and no survival benefits of group membership, the reduced incentives to remain in the natal group drive most subordinates to disperse and breed elsewhere (average dispersal rate between 88% in harsh environments and 93% in benign environments; Table S1). Adding a reaction norm of dispersal to dominance rank did not change the results (see the section “Dominance-dependent dispersal propensities” and Figure S5 in the Supplemental Material). The few subordinates that remain in the natal territory help to obtain indirect benefits (average within-group relatedness between 0.36 in benign environments and 0.72 in harsh environments; Table S1). Even though individuals in our model cannot adjust their help according to their relatedness to the breeder, the high levels of relatedness that evolved under this paradigm indicate that allowing for targeted help would not qualitatively change the results. Survival costs of defensive tasks are then avoided, as mortality without group benefits is relatively high, while costs to dominance value are small for close relatives because of indirect fitness benefits and low conflict of interest (Figure 2; KS). Adjusting input parameters to obtain higher survival probabilities and favoring philopatry did not qualitatively change the results, except for extreme values of baseline survival (see Supplementary Material). Only when (1) there is a tenfold increase in the productivity of the breeders per unit of help, and (2) division of labor is required for a productivity boost, do group members stay philopatric and coordinate labor to help raise kin in some circumstances (Figure S4). Hence, kin selection alone is unlikely to select for the evolution of defensive tasks and division of labor in vertebrates.

Revised version: In contrast to expectations from models of eusocial insects, our vertebrate model does not readily evolve division of labor when only kin selection is considered (see the section “Kin selection and the evolution of division of labor” and Figures S1-S4 in the Supplemental Material for more details on further exploration on conditions that may favor the evolution of division of labor under only kin selection). Specifically, forms of help that impact survival rarely evolve under any environmental condition when only kin selection occurs (Figure 2; KS). This occurs because with kin competition for the breeding position and no survival benefits of group membership, the reduced incentives to remain in the natal group drive most subordinates to disperse and breed elsewhere (average dispersal rate between 88% in harsh environments and 93% in benign environments; Table S1). Adding a reaction norm of dispersal to dominance rank did not change the results (see the section “Dominance-dependent dispersal propensities” and Figure S5 in the Supplemental Material). The few subordinates that remain in the natal territory help to obtain indirect benefits (average within-group relatedness between 0.36 in benign environments and 0.72 in harsh environments; Table S1). Even though individuals in our model cannot adjust their help according to their relatedness to the breeder, the high levels of relatedness that evolved under this paradigm indicate that allowing for targeted help would not qualitatively change the results. Survival costs of defensive tasks are then avoided, as mortality without group benefits is relatively high, while costs to dominance value are small for close relatives because of indirect fitness benefits and low conflict of interest (Figure 2; KS). Adjusting input parameters to reduce defense costs on survival or to increase survival probabilities and favor philopatry did not qualitatively change the results, except for extreme values of baseline survival (see Supplementary Material; Figure S4 & S7). Only when (1) there is a tenfold increase in the productivity of the breeders per unit of help, and (2) division of labor is required for a productivity boost, do group members stay philopatric and coordinate labor to help raise kin in some circumstances (Figure S4). Hence, kin selection alone is unlikely to select for the evolution of defensive tasks and division of labor in vertebrates.

Earlier version: Conversely, under only direct fitness benefits (group augmentation), tasks associated with costs to the dominance value are costlier because the only path to obtain fitness benefits is via direct reproduction. Therefore, defensive tasks that do not affect the dominance value are favored (Figure 2; GA). Under this scenario, division of labor can evolve to maximize the reproductive output of the breeders (Figure 2; GA, m = 0.3, xh ≤ 5), but also because it is the optimal strategy (Figure 2; GA, m = 0.3, xh = 7). In benign environments, however, no help evolves due to the lower incentives to increase survival and the steep competition for reproduction resulting from habitat saturation.

Revised version: Conversely, under only direct fitness benefits (group augmentation), tasks associated with costs to the dominance value are costlier because the only path to obtain fitness benefits is via direct reproduction. Therefore, defensive tasks that do not affect the dominance value are favored (Figures 2 & S7; GA). Under this scenario, division of labor can evolve to maximize the reproductive output of the breeders (Figure 2; GA, m = 0.3, xh ≤ 5), but also because it is the optimal strategy (Figure 2; GA, m = 0.3, xh = 7). In benign environments, however, no help evolves due to the lower incentives to increase survival and the steep competition for reproduction resulting from habitat saturation.

Earlier version: When subordinates can obtain both direct and indirect fitness, division of labor is more strongly selected in increasingly harsher environments to increase breeder productivity (Figure 2; KS + GA). Overall, in moderate and harsh environments, work tasks with costs to the dominance value are more strongly selected than in benign environments due to high mortality increasing both the breeder’s turnover, reducing breeding competition (lower effective costs of work tasks), and the costs to survival (high costs of defensive tasks). However, in more benign (and often highly productive) environments that lead to habitat saturation, help likely evolved initially in family groups, and defensive tasks are favored because competition for the breeding position is lower under kin selection. In this scenario, kin selection acts not only by favoring subordinates in their natal group to increase the productivity of a related breeder (Figure 2; KS), but also by the increased survival of their siblings when augmenting the group size (Figure 2; KS + GA), regardless of whether they are related to the current breeder.

Revised version: When subordinates can obtain both direct and indirect fitness, division of labor is more strongly selected in increasingly harsher environments to increase breeder productivity, a result mainly driven by an increase in direct fitness benefits (Figure 2; KS + GA). Conversely, when direct fitness benefits are low, division of labor primarily emerges in benign environments from a combination of both direct and indirect fitness benefits (Figure S7; KS + GA). Additionally, in moderate and harsh environments, work tasks with costs to the dominance value are more strongly selected than in benign environments due to high mortality increasing the breeder’s turnover, which reduces breeding competition (lower effective costs of work tasks), and the costs to survival (high costs of defensive tasks), while defense tasks are mainly promoted by direct fitness benefits and in more benign conditions (Figure S7). Hence, in more benign (and often highly productive) environments that lead to habitat saturation, help likely evolved initially in family groups, and defensive tasks are favored because competition for the breeding position is lower under kin selection. In this scenario, kin selection acts not only by favoring subordinates in their natal group to increase the productivity of a related breeder (Figure 2; KS), but also by the increased survival of their siblings when augmenting the group size (Figure 2; KS + GA), regardless of whether they are related to the current breeder. In harsher environments, however, competition between unrelated individuals decreases in favor of coordinated cooperation to reduce mortality (Figures 2 & S7; GA). Together, these patterns indicate that direct fitness benefits are the primary force driving cooperative specialization in harsh environments, where increased breeder productivity boosts the survival of both helpers and their relatives (Figure 2). In contrast, in benign environments, cooperation arises mainly through kin-selected benefits, with subordinates contributing primarily by performing their preferred task of defending the breeder’s offspring (Figure 2).

Removed in the revised version: Overall, division of labor seems to be more likely to evolve in harsh environments where raising breeder productivity effectively enhances the survival of both helpers and relatives (Figure 2). In contrast, help evolves in benign environments largely due to kin selection benefits and subordinates performing their preferred task of defending the breeder’s offspring (Figure 2).

The role of dominance in task specialization

Earlier version: In the previous section, we explored the role of direct and indirect fitness benefits in the evolution of division of labor via task specialization. However, determining whether parameter γR (i.e., the direction and rate of change in task specialization with dominance) significantly differs from zero is crucial. Intermediate values of phenotypic task specialization may indicate division of labor among different age-or dominant-related groups (if γR ≠ 0), or lack of specialization in either task (if γR = 0). Hence, for age-polyethism or temporal division of labor to arise in our model, individuals must switch between tasks based on their age-related dominance (probability of choosing a “defensive task” versus a “work task” T ≠ 0 or 1, and γR ≠ 0; Table 1). Across all scenarios, individuals are selected to perform tasks based on their probability of inheriting the territory (Figure 3A). Specifically, when subordinates prefer defensive tasks (Figure 3A, m < 0.2, circles), they engage in work tasks at a young age and increasingly undertake defensive tasks as their dominance rises (Figure 3, m < 0.2). Similarly, if their preference leans toward work tasks (Figure 3A, m ≥ 0.2, circles), younger subordinate helpers and those of lower dominance focus more on defense initially, but switch to work tasks as they approach the breeding stage (Figure 3, m ≥ 0.2). In essence, subordinate helpers allocate more time to tasks with lower fitness costs as they near a competitive dominance value, all without any active enforcement by other group members, resulting in age-or rank-polyethism. These results are consistent with the rise in division of labor, defined as within-individual consistency and between-individual differences in task choice [64].

Revised version: In the previous section, we explored the role of direct and indirect fitness benefits in the evolution of division of labor via task specialization. However, determining whether parameter γR (i.e., the direction and rate of change in task specialization with dominance) significantly differs from zero is crucial. Intermediate values of phenotypic task specialization may indicate division of labor among different age– or dominant-related groups (if γR ≠ 0), or lack of specialization in either task (if γR = 0). Hence, for age-polyethism or temporal division of labor to arise in our model, individuals must switch between tasks based on their age-related dominance (probability of choosing a “defensive task” versus a “work task” T ≠ 0 or 1, and γR

Added in the revised version: ≠ 0; Table 1). Across all scenarios, individuals are selected to perform tasks based on their probability of inheriting the territory (Figure 3A). Specifically, when subordinates prefer defensive tasks (Figure 3A, m < 0.2, circles), they engage in work tasks at a young age and increasingly undertake defensive tasks as their dominance rises (Figure 3, m < 0.2). Similarly, if their preference leans toward work tasks (Figure 3A, m ≥ 0.2, circles), younger subordinate helpers and those of lower dominance focus more on defense initially, but switch to work tasks as they approach the breeding stage (Figure 3, m ≥ 0.2). In essence, subordinate helpers allocate more time to tasks with lower fitness costs as they near a competitive dominance value, all without any active enforcement by other group members, resulting in age– or rank-polyethism. These results are consistent with the rise in division of labor, defined as within-individual consistency and between-individual differences in task choice [64].

Evolved reaction norms to age on the display of task specialization.

Earlier version: The evolutionary equilibria for the reaction norms of task specialization are shown at five different levels of environmental quality, ranging from benign (m=0.1, purple) to harsh (m=0.3, yellow), and xh=5. A: γR > 0 signifies that individuals increase the probability of performing work tasks with dominance (defense ➔ work), whereas γR < 0 signifies that individuals increase the probability of performing defensive tasks with dominance (work ➔ defense). The optimum breeder productivity per unit of help provided was either when both tasks were performed to a similar extent, potentially selecting for division of labor (▴), or when no restrictions were introduced to the task performed by the group members (〇). Results show that at equilibrium when division of labor evolves, individuals increase the probability of performing their preferred task (circles), when increasing their dominance value. B: Evolved reaction norms to dominance value with average γ0 and γRacross 20 replicas for varying quality environments. All parameter values described in Table 1.

Revised version: The evolutionary equilibria for the reaction norms of task specialization are shown at five different levels of environmental quality, ranging from benign (m=0.1, purple) to harsh (m=0.3, yellow), and xh=5. A: γR > 0 signifies that individuals increase the probability of performing work tasks with dominance (defense➜ work), whereas γR < 0 signifies that individuals increase the probability of performing defensive tasks with dominance (work ➜ defense). The optimum breeder productivity per unit of help provided was either when both tasks were performed to a similar extent, potentially selecting for division of labor (▴), or when no restrictions were introduced to the task performed by the group members (〇). Results show that at equilibrium when division of labor evolves, individuals increase the probability of performing their preferred task (circles), when increasing their dominance value. B: Evolved reaction norms to dominance value with average γ0 and γR across 20 replicas for varying quality environments. All parameter values described in Table 1.

Discussion

Our findings suggest that the direct benefits of group living play a driving role in the evolution of division of labor via task specialization in species with totipotent workers, including all vertebrates and a few insects. In contrast, in eusocial species characterized by high relatedness and permanent worker sterility, such as in most eusocial insects, workers acquire fitness benefits only indirectly, and so, their fitness is contingent upon the overall success of the colony rather than the individual [65, 66]. Group members in such eusocial species are therefore predicted to maximize colony fitness due to the associated lower within-group conflict, and to favor the evolution of division of labor to increase group productivity efficiency [51, 52, 54]. However, in taxa with lower within-group relatedness, such as nearly all vertebrates, division of labor is facilitated when cooperation is crucial for survival and when the benefits of cooperation are evenly distributed among group members [3]. In particular, survival fitness benefits derived from living in larger groups seem to be key for the evolution of cooperative behavior in vertebrates [25, 67], and may also translate into low within-group conflict. This suggests that selection for division of labor in vertebrates is stronger in smaller groups, where increased productivity significantly enhances survival, as opposed to research on eusocial insects that posits division of labor is more likely to evolve in larger groups [68–70]. That is not to say, however, that kin selection is unimportant in cooperatively breeding vertebrates or insects. Our model demonstrates that cooperative breeders are more likely to evolve division of labor when direct and indirect fitness benefits act in concert. Moreover, our model predicts that the type of task that individuals prefer depends on the type of fitness benefit that subordinates can attain. While direct fitness benefits in the form of group augmentation select more strongly for defensive tasks, kin selection alone seems to select only for work tasks. In groups of unrelated individuals where members can only gain direct fitness benefits, costs to dominance value and to the probability of attaining a breeding position are comparatively larger than survival costs. Conversely, groups with no group benefits tend to be small and suffer high mortality, for which defensive tasks only aggravate the situation.

Previous models of task allocation in eusocial insects have focused on how seemingly identical individuals of the same age and broad morphology may specialize in different types of tasks [8, 64]. In contrast, our model shows that extrinsic ecological factors are also important for the evolution of task specialization in addition to intrinsic factors such as age or dominance rank. In short, our model predicts that task specialization, which varies with age or rank, will depend on environmental quality. In more favorable environments, helpers are predicted to increase defensive tasks with age or rank, whereas in harsh environments, work tasks are predicted to increase with age or rank. Furthermore, harsher environments exert stronger selective pressure for task specialization due to increased survival challenges. Enhanced cooperation and improved group productivity through helping mitigate the negative effects of the environment. This finding is consistent with previous research on the role of ecological constraints in driving the evolution of cooperation [71–73].

Although rare, evidence of division of labor in cooperative breeding vertebrates comes from a few studies in mammals [74], birds [75, 76], and fishes [77, 78]. Our model predictions are broadly consistent with empirical data from many of these studies. For example, support for the role of environmental harshness on the evolution of division of labor via age-dependent task specialization in vertebrates comes from arguably the most similar vertebrate system to eusocial insects, the naked mole-rat (Heterocephalus glaber). Naked mole-rats live in self-dug burrows that protect their inhabitants from predation and climatic extremes [74]. Therefore, even though they inhabit habitats with scarce and unpredictable rainfall, environmental harshness is likely to be limited in their stable subterranean burrows, with mortality mainly spiking during dispersal attempts. Model predictions for species that evolved in benign environments are then met, with individuals shifting from work tasks such as foraging for food, digging, and maintaining the burrow system, to defensive tasks such as guarding and patrolling as individuals grow older and larger [74, 79, 80]. Although naked mole-rats exhibit high levels of within-colony relatedness [81], group size also exerts a positive influence on survival, potentially generating group augmentation benefits [82, 83]. Under these conditions, our model predicts the highest levels of task partitioning and division of labor. In contrast, Neolamprologus pulcher, a cooperatively breeding cichlid fish that lives in harsh environments with high risk of predation [62], shows the opposite tendency. In line with our model predictions, larger and older helpers of this species invest relatively more in territory maintenance, whereas younger/smaller helpers defend the breeding shelter of the dominant pair to a greater extent against experimentally exposed egg predators [78]. Territory maintenance has been shown to greatly affect routine metabolic rates and, hence, growth rates [84], which directly translates into a decrease in the probability of becoming dominant and attaining breeding status. In addition, both group augmentation and kin selection have been found to play a role in the evolution of help in this species [23, 55]. More research is needed to verify whether this trend is widespread in cooperatively breeding vertebrates with division of labor. The preferred form of help provided by individuals throughout their life course might deviate from our expectations if helpers exhibit differing abilities to perform distinct tasks at different ages or life stages. We assumed that individuals of all ages possessed equal capability to perform different tasks. However, as individuals age and their size and body mass increase, they may become more adept at certain tasks such as deterring predators [36, 85–89]. Body condition that is unrelated to age may also influence helper contributions, as individuals in poorer condition may be more constrained on their energy allocation in activities other than self-maintenance [90–92]. Furthermore, our model does not capture instances in which help in general or specific tasks are enforced by the dominant pair or other more dominant group members [93–95].

Despite evidence of task specialization and division of labor across several vertebrate taxa, other empirical studies have failed to find similar results. For instance, researchers failed to find evidence of task specialization in Damaraland mole-rats (Fukomys damarensis), despite their social structure resembling that of naked mole-rats [96, 97], with some variation in digging behavior with age and sex [98]. However, it is worth noting that researchers have only investigated correlates between different work tasks, potentially overlooking negative correlations with other helping activities that pose a more direct cost to survival. In fact, many studies that reported an absence of division of labor did not consider defensive tasks in their analyses [96, 97, 99, 100]. Nonetheless, studies exist in which researchers did not find a negative correlation between defense and offspring provisioning or territory maintenance [42, 101, 102]. However, individuals may exhibit a positive or non-relationship between tasks if they also vary in total helpfulness across their lifetime, but may still switch the proportion of time/effort allocated to different types of tasks. In our model, we addressed for this potential tendency by restricting individuals from changing their overall helpfulness with age/dominance rank. We encourage future researchers to consider the proportion of time/energy investment in addition to examining correlations between tasks.

In summary, our study helps to elucidate the potential mechanisms underlying division of labor between temporal non-reproductives via task specialization in taxa beyond eusocial organisms. Harsh environments, where individuals can obtain direct fitness benefits from group living, favor division of labor thereby enhancing group productivity and, consequently, group size. Variation in the relative fitness costs of different helping tasks with age favors temporal polyethism, with individuals changing their behavior patterns to improve their chances of becoming breeders. Hence, future empirical research on division of labor should prioritize cooperatively breeding species inhabiting relatively harsh environments where subordinates benefit greatly from group membership, such as protection or improved food acquisition. Additionally, we emphasize the importance of including sentinel and defense behaviors in future studies to encompass a broader spectrum of fitness costs of cooperation, and to investigate potential changes in task investment throughout the helpers’ lifetimes.

Materials and Methods

We developed an individual-based model in which individuals may display and vary their efforts in different helping tasks. The ancestral state features an absence of alloparental care and no task specialization. We assume that residing in larger groups yields overall survival advantages, potentially resulting in direct fitness benefits associated with the increase in group size, as proposed by the group augmentation hypothesis. We also include the coevolution of the helping strategies with their dispersal propensity, as dispersal affects group size and kin structure, and thus may impact the strength of indirect fitness benefits.

Breeding cycle

We model a population consisting of a fixed number of breeding territories, each comprising a dominant breeder that reproduces asexually and monopolizes reproduction, as well as an undefined number of subordinates that queue for the breeding position. The number of subordinates is determined by the productivity of the breeders, the dispersal decisions of the group members, and immigration. Subordinates may help increase the fecundity of the breeder, so that the breeder’s fecundity depends on the cumulative level of help provided by the subordinates within the group, with an optimum increase when different helping tasks are performed to a similar degree.

Offspring inherit the genes for dispersal and helping tendencies from the breeders, potentially exhibiting slight variation in the presence of mutations (Figure 1, Step 1). Each genetic trait is controlled by a single locus, and may take any real number. Mutations occur at each genetic locus at a frequency of μ, causing minor adjustments to the allele’s value inherited from the parent. This alteration involves the addition of a value drawn from a normal distribution characterized by a mean of 0 and a standard deviation of σμ (Table 1).

After breeders reproduce, their offspring and the subordinates in the group may disperse independently from each other. Dispersal propensity is controlled by a gene with values ranging from 0 to 1, where 0 represents certain philopatry and 1 represents absolute dispersal (Table 1). In each breeding cycle, dispersers may migrate to another group to become subordinates or remain as floaters waiting for breeding opportunities, which is also controlled by the same genetic dispersal predisposition as subordinates (Figure 1, Step 2). Successful dispersers enter the breeding queue within the group based on their dominance value rather than starting at the bottom of the hierarchy. Subordinate group members, either natal individuals born in the group or immigrants to the group, express some level of help in the form of alloparental care, which can potentially evolve to zero (Figure 1, Step 3). To be conservative, we use the term help as synonymous with altruistic behavior, which increases the fitness of the beneficiary (i.e., breeder) with an immediate decrease in the actor’s fitness. Hence, cooperative behaviors such as defense from adult predators that incur immediate benefits to the actor as well as by-product mutualistic benefits to other group members, are not considered in this study. Dispersal propensity, alloparental care provisions, and the type of helping task they perform all depend on their genetic predisposition. They may perform tasks that reduce their body condition and, therefore, their chances of becoming breeders, such as feeding offspring or territory maintenance (hereafter called “work tasks”), or they may perform tasks that reduce their immediate survival such as predator defense (hereafter called “defensive tasks”) [35]. For simplification, we assume complete dichotomous, short-term non-cumulative fitness consequences between the different types of tasks (i.e., survival and dominance costs are only sustained during the breeding season in which they perform the task, and individuals only perform one of the tasks in the given cycle). We note that the names given for each type of task are meant to facilitate interpretation, and that in some species, some tasks classified as work tasks may have a higher immediate cost in survival than in the probability to breed, and vice versa. However, helping tasks should be classified in empirical systems according to the primary fitness cost associated with each type of helping activity (e.g., if a work activity mostly negatively impacts survival in a given species, this task is described in this paper as defensive).

Individual survival depends on environmental quality, group membership, group size, and the amount of help provided during defensive tasks to the breeder’s offspring (Figure 1, Step 4). Environmental quality is expressed as a varying maximum survival probability, independent of social factors. Additional environmental effects, such as variation in mortality linked to dispersal, or on the probability of floaters finding a new group to start breeding, have been explored elsewhere [67].

If the breeder in a group dies, all subordinates in the group as well as a randomly drawn sample of floaters from outside the group compete for the breeding position (Figure 1, Step 5). The number of floaters bidding for an empty breeding spot is then proportional to the relative abundance of floaters with regard to the total number of breeding territories, to account for spatial viscosity and the probability of being accepted in a new group (Table 1). In this model, age influences the probability of attaining the breeding position. Older individuals, owing to their higher resource holding potential, have a greater probability of becoming breeders, as observed in various species [62, 103–111]. Age in our model corresponds to the number of breeding cycles or seasons that an individual survives. Additionally, work tasks reduce an individual’s dominance value as a result of the impact on their body condition. The probability of filling an empty breeding position is then implemented as a lottery weighted by the dominance value of the candidates. Therefore, subordinates may stay in the natal group and queue to inherit the dominant position [105, 112–115], but they may also disperse and breed or queue in another group [37, 113, 116]. We ran simulations for 200’000 breeding cycles or until equilibrium was reached for all genetic traits, across 20 replicas to assess consistency.

Age-dependent task specialization

Earlier version: The probability of choosing defensive versus work tasks T takes a logistic function with boundaries between 0 and 1 as given in Eq. 1, in which individuals may adjust their helping task as they age and become more dominant, where γRmodifies the strength and direction of the effect of dominance (R) on task choice, and γ0acts as the intercept. Therefore, when T approaches zero, individuals specialize in work tasks throughout their lives, and when T approaches one, they specialize in defensive tasks. If γR ≠ 0, given that T ≠ 0 and T ≠ 1, individuals show age-dependent task specialization.

Revised version: The probability of choosing defensive versus work tasks T takes a logistic function with boundaries between 0 and 1 as given in Eq. 1, in which individuals may adjust their helping task as they age and become more dominant, where γR modifies the strength and direction of the effect of dominance (R) on task choice, and γ0 acts as the intercept. Therefore, when T approaches zero, individuals specialize in work tasks throughout their lives, and when T approaches one, they specialize in defensive tasks. If γR ≠ 0, given that T ≠ 0 and T ≠ 1, individuals show age-dependent task specialization.

In addition, the phenotypic expression of help (H) is regulated by another gene and equals the allelic value of gene α, which remains fixed throughout the helper’s life. Even though in nature individuals are likely to change the total amount of help given throughout their lives, we make this simplification to allow easier interpretation of how individuals may switch between task types as they age. How individuals in the model may adapt their level of help with age and social and environmental conditions has been described elsewhere [67]. Hence, the gene α regulates the amount of help expressed, while the genes γ determine which specific helping tasks are performed at different time points in the breeding cycle. The fitness costs of helping are described in the following section.

Fitness costs of different helping tasks

Helping tasks are classified in our model according to their short-term fitness consequences. We assume that work tasks incur a more immediate cost to the body condition of helpers because of expenses in terms of time and energy. Work tasks reduce the dominance value R to inherit the breeding position in our model following Eq. 2, where γh is a parameter that influences the strength of the cost of help (Hwork) in the dominance value. We assume that dominance is influenced by age (t) as a proxy for resource holding potential, with older individuals having a greater dominance value, Note that the cost to an individual’s dominance value is influenced by their current age and the current specific task that they are performing, the cost of which is not cumulative over time. Therefore, an individual’s effective rank within the dominance hierarchy is, irrespective of whether they belong to a group or not, determined by their dominance value relative to other competitors for the breeding position, with higher dominance values corresponding to higher ranks in the hierarchy.

In contrast, we assume that defensive tasks have a more immediate fitness cost on survival (e.g., predator or space competitor defense) due to the high risk of injury. Therefore, subordinate helpers that defend have a survival probability (SH) that decreases with the amount of help provided (Hdefense) given by the logistic Eq. 3.1, where xh is a parameter that influences the strength of the cost of help in survival, In addition to defensive tasks, social and ecological factors also influence the survival of group members. The term m denotes the baseline mortality that determines how harsh the environment is. Higher values of m indicate higher overall baseline mortality for all individuals in the population, irrespective of social factors. For simplicity, we assume that the harshness of the environment only impacts survival, although other effects such as on reproductive output or the strength in which help improves fecundity are likely to have an effect in nature [117]. Group size (N) provides survival benefits, such as those offered by safety in numbers or increased resource defense potential, for both subordinates and breeders (Eq. 3.1 and 3.2 respectively), where xn is a parameter that quantifies the effect size of the benefit of group size in survival, and x0 an intercept. Floaters sustain the highest mortality (SF, Eq. 3.3) since they do not benefit from group protection,

Need for division of labor

To assess the rules governing task specialization and division of labor, we first outline a basic model where individuals evolve their preferred helping task. In this model, alloparental help increases breeder productivity regardless of how tasks are distributed among helpers. We then extend this framework to a second model in which the breeder’s reproductive outcome is maximized only when the group’s efforts are roughly balanced across the two types of tasks, thereby actively promoting division of labor. Thus, in the second scenario, the maximum effective contribution of cumulative group help of each task type (Hmax) that can enhance fecundity is given by Eq. 4: where km is a parameter that relaxes the requirement for perfectly balanced task allocation, allowing fecundity to increase even if the contributions to defense and work are somehow uneven. The breeders then reproduce asexually. The total number of offspring produced (K) is contingent upon the breeders’ baseline fecundity (k0) and the cumulative helping effort of all subordinates. Following Eq. 5, fecundity increases with total help but exhibits diminishing returns, where kh is a parameter that scales the benefit of cumulative help on fecundity. To prevent exceeding the maximum effective help, any cumulative defense or work effort that surpasses Hmax is truncated at Hmax (i.e. if ). The breeder’s fecundity is then determined by a randomly selected value drawn from a Poisson distribution with mean of K.

Fitness benefits of helping

We include both direct and indirect fitness benefits of helping to assess the relative importance of each in the evolution of task specialization. Indirect fitness benefits from relatedness among group members (i.e., kin selection) emerge from demographic factors such as dispersal tendencies, mortality rates, and group dynamics, though they were not coded explicitly. To calculate the coefficient of relatedness between the breeder and subordinates in a group, we calculated the coefficient of a linear regression between the allelic values of the breeders and helpers of a neutral gene that changes exclusively by genetic drift [118, 119]. Direct benefits may also be obtained by enhancing group size by helping to increase the breeder’s fecundity, as larger groups benefit from higher survival probabilities in our model (i.e., group augmentation).

To distinguish the effect of kin selection from group augmentation, we created two parallel models for comparison to the default model (group augmentation and kin selection acting independently), in which one of the fitness benefits was removed. In the “only group augmentation” implementation, individuals just born (i.e., age = 1) that decide to stay in the natal territory as subordinates (Figure 1, Step 2) are shuffled into another group, thereby removing the influence of relatedness from the model. In the “only kin selection” implementation, we removed the survival benefits linked to living in a group (xn = 0).